<a href="https://colab.research.google.com/github/ddaeducation/DataAnalyst/blob/main/Practice_Activity_Using_Dictionaries.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Activity: Using Dictionaries

## Introduction

In this activity you will practice the following:
- Creating dictionaries
- Accessing values in a dictionary
- Adding key-value pairs to a dictionary
- Modifying values in a dictionary
- Dictionary methods


#### Question 1

Create a dictionary called `student` which contains the following:
- name: derrick
- age: 17
- math_grade: 90


<!-- grade_id: cell-f574fadfa206fd79, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-290f7938a07cd6d9, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(student, dict), 'Have you created a dictionary called student?'
assert len(student) == 3, 'Have you added the correct number of key-value pairs to the dictionary?'

```

#### Question 2

Access the value of the `math_grade` key in the `student` dictionary you created in Question 1 and assign it to a variable called `grade`.


<!-- grade_id: cell-92771faa1a1ba3f7, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-c024fb20adf99fd1, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(grade, int), 'Have you accessed the value of the math_grade key in the student dictionary?'

```

#### Question 3

Add a new key-value pair to the `student` dictionary you created in Question 1. The new key should be `english_grade` and the value should be `85`.


<!-- grade_id: cell-28dfce50b2efee61, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-ad33c27005dd25c1, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert len(student) == 4, 'Have you added the correct number of key-value pairs to the dictionary?'

```

#### Question 4

Using the given dictionary `top_student`, store all the values in the dictionary in a variable called `top_student_grades`.


```python
top_student = {
    "math_grade": 95,
    "english_grade": 98,
    "history_grade": 90,
    "science_grade": 93,
    "art_grade": 92,
    "music_grade": 96
}
```

<!-- grade_id: cell-1d292257a672451a, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-d2323c8fa84b0fc1, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks


```