<a href="https://colab.research.google.com/github/ddaeducation/DataAnalyst/blob/main/Using_booleans.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Activity: Using Booleans

## Introduction

In this activ|ity, you will use booleans to complete the following questions.
This activity includes:
- Logical operators
- Boolean expressions
- Membership operators


#### Question 1

Check if the variable age is greater than or equal to 18 and less than or equal to 25. Assign the result to a variable called is_age_valid.

If the age is valid, negate the variable `can_drive` and assign the result to a new variable called `cannot_drive`.


```python
can_drive = False
age = 20
```

<!-- grade_id: cell-9640f4cde71b29f1, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-887f1340ad039e18, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(cannot_drive, bool), "Make sure you are assigning a boolean value to cannot_drive"
assert isinstance(is_age_valid, bool), "Make sure you are assigning a boolean value to is_age_valid"

```

#### Question 2

Check if the variable `name` is in the list `names`. Assign the result to a variable called `is_name_valid`.


```python
name = "John"
names = ["John", "Jane", "Jack"]
```

<!-- grade_id: cell-b04289a88517ce29, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-8df0006d2710420b, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(is_name_valid, bool), "Make sure you are assigning a boolean value to is_name_valid"

```

#### Question 3

Write a boolean expression that checks if the variable `whole_num` is not equal to `float_num`. Assign the result to a variable called `is_not_equal`.


```python
whole_num = 5
float_num = 5.0
```

<!-- grade_id: cell-3764d95214f9d8e6, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-b89d59ea6b9ee8dd, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(is_not_equal, bool), "Make sure you are assigning a boolean value to is_not_equal"

```