<a href="https://colab.research.google.com/github/ddaeducation/DataAnalyst/blob/main/V2_using_lists_and_tuples.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Activity: Using Lists & Tuples

## Introduction

In this activity you will practice using and creating lists and tuples.
This activity contains the following topics:
- Creating lists
- Creating tuples
- Indexing
- Slicing
- List methods
- Unpacking


#### Question 1

Create a list called `countries` with the names of 5 different countries of your choice.


<!-- grade_id: cell-abc66a5348e14a02, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here

```

<!-- grade_id: cell-2ce11156bb8d81a2, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(countries, list), 'Your variable is not a list'

```

#### Question 2

Using the `countries` list, assign the first and last countries in the list to the variables `first_country` and `last_country` respectively.


<!-- grade_id: cell-f25ba6392107e856, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-1a0e87f96e13f6a4, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(first_country, str), 'Your variable is not a string'
assert isinstance(last_country, str), 'Your variable is not a string'

```

#### Question 3

Using the `countries` list, add a country to the end of the list and then assign the last 3 countries to a variable called `last_three_countries`.


<!-- grade_id: cell-5bc224fb57e5e3f0, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-8fc6b48f4b0c446a, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(last_three_countries, list), 'Your variable is not a list'

```

#### Question 4

You're given a tuple containing superheroes, unpack the tuple into variables `superman`, `batman`, and `black_panther`


```python
# Given tuple

superheroes = ('Superman', 'Batman', 'Black Panther')
```

<!-- grade_id: cell-b1e856700a857c95, points: 0, grade: false, solution: true -->
```python
# Your code here

# your code here
```

<!-- grade_id: cell-daf55b2d5796fc09, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert isinstance(superman, str), 'Your variable is not a string'
assert isinstance(batman, str), 'Your variable is not a string'
assert isinstance(black_panther, str), 'Your variable is not a string'

```